<!-- grade_id: cell-4cb12024ca3ad2c2, points: 1, grade: true, solution: true -->
```scala
def add(x: Int, y: Int): Int = {
    // BEGIN SOLUTION
    x + y
    // END SOLUTION
}
```

Output:
```
defined [32mfunction[39m [36madd[39m
```

<!-- grade_id: cell-e580d9d8cfd4f99e, points: 1, grade: true, solution: false -->
```scala
assert(add(4, 5) == 9)
assert(add(4, -1) == 3)
// BEGIN HIDDEN TESTS
assert(add(9, 10) == 19)
// END HIDDEN TESTS
```

```scala

```